---
title: "Bike Rentals in D.C."
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
    css: "css/font-size.css"
runtime: shiny_prerendered
description: >
  Learn about modelling, interpretations of model performance, and backward selection
  by analysing bike share data from Washington D.C.
---



```{r setup, include=FALSE}
# load packages ----------------------------------------------------------------

library(tidyverse)
library(tidymodels)
library(dsbox)
library(learnr)
library(gradethis)
library(learnrhash)
library(GGally)

# set options for exercises and checking ---------------------------------------
gradethis_setup()

# hide non-exercise code chunks ------------------------------------------------
knitr::opts_chunk$set(echo = FALSE)
# data prep --------------------------------------------------------------------
# create updted bike dataset for all exercises after wrangling
mutatedbike <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ),
    season = fct_relevel(season, "spring")) %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    holiday = fct_relevel(holiday, "no"),    
    workingday = ifelse(workingday == 0, "no", "yes"),
    workingday = fct_relevel(workingday, "no")
  ) %>%
  mutate(
    yr = ifelse(yr == 0, "2011", "2012"),
    yr = fct_relevel(yr, "2011")
  ) %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      weathersit == 2 ~ "mist",
      weathersit == 3 ~ "light precipitation",
      weathersit == 4 ~ "heavy precipitation"
    ),
    weathersit = fct_relevel(weathersit, "clear")
  ) %>%
  mutate(
    temperature_raw = temp * 41,
    feeling_temperature_raw = atemp * 50,
    humidity_raw = hum * 100,
    windspeed_raw = windspeed * 67
  )
```







## Introduction

![Photo by El Salanzo on Unsplash](images/el-salanzo-yUvQnZY29Oc-unsplash.jpg)


Bike sharing systems are a new generation of traditional bike rentals where the whole process from membership, rental and return back has become automatic. 
Through these systems, the user is able to easily rent a bike from a particular position and return it back at another position. 
As of May 2018, there are about over 1600 bike-sharing programs around the world, providing more than 18 million bicycles for public use.
Today, there exists great interest in these systems due to their important role in traffic, environmental and health issues. 

Apart from interesting real world applications of bike sharing systems, the characteristics of data being generated by these systems make them attractive for the research. 
Opposed to other transport services such as bus or subway, the duration of travel, departure and arrival position is explicitly recorded in these systems. 
This feature turns bike sharing system into a virtual sensor network that can be used for sensing mobility in the city. 
Hence, it is expected that most of important events in the city could be detected via monitoring these data.

Source: [UCI Machine Learning Repository - Bike Sharing Dataset](http://archive.ics.uci.edu/ml/datasets/Bike+Sharing+Dataset)

### Learning goals

- Continue to hone your data wrangling skills.
- Practice modelling and interpreting model results and performance. 
- Conduct backwards selection for finding the "best" model. 

### Packages

Load the packages we need for this tutorial using the chunk below.

```{r load-packages, exercise = TRUE}
library(tidyverse)
library(broom)
library(dsbox)
```

```{r load-packages-solution}
library(tidyverse)
library(broom)
library(dsbox)
```

```{r load-packages-check}
grade_this_code("The packages are now loaded!")
```

## Data and codebook

The data include daily bike rental counts (by members and casual users) of Capital Bikeshare in Washington, DC in 2011 and 2012 as well as weather information on these days. 
The data is contained in the `dsbox` package and is called `dcbikeshare`.

Go ahead and `glimpse` the data to take a peek.

```{r glimpse-data-setup}

dcbikeshare <- dsbox::dcbikeshare

```

```{r glimpse-data, exercise=TRUE, exercise.setup="glimpse-data-setup"}


```


A description of the variables is below:

| Variable name    | Description 
|:--------|:-------------------------------------------------------------
| `instant`		| record index
| `dteday` 		| date
| `season` 		| season (1:winter, 2:spring, 3:summer, 4:fall)
| `yr` 		    | year (0: 2011, 1:2012)
| `mnth` 		  | month (1 to 12)
| `holiday` 	| whether day is holiday or not (extracted from http://dchr.dc.gov/page/holiday-schedule)
| `weekday` 	| day of the week
| `workingday`| if day is neither weekend nor holiday is 1, otherwise is 0.
| `weathersit`| 1: Clear, Few clouds, Partly cloudy, Partly cloudy
|             | 2: Mist + Cloudy, Mist + Broken clouds, Mist + Few clouds, Mist
|             | 3: Light Snow, Light Rain + Thunderstorm + Scattered clouds, Light Rain + Scattered clouds
|             | 4: Heavy Rain + Ice Pallets + Thunderstorm + Mist, Snow + Fog
| `temp` 			| Normalized temperature in Celsius. The values are divided by 41 (max)
| `atemp`			| Normalized feeling temperature in Celsius. The values are divided by 50 (max)
| `hum` 			| Normalized humidity. The values are divided by 100 (max)
| `windspeed`	| Normalized wind speed. The values are divided by 67 (max)
| `casual`		| Count of casual users
| `registered`| Count of registered users
| `cnt` 		  | Count of total rental bikes including both casual and registered


In this tutorial, we will consider our response variable to be `cnt`, the count of total rental bikes. The explanatory variables are several some of which are categorical (`season`, `yr`, `holiday`, `workingday`, weathersit`) and others numeric. 

## Recoding data

In the `dcbikeshare` data frame, all the variables have been entered as numeric values regardless of their type. We will therefore begin by recoding many of the categorical explanatory variables. 

### Recode season

As noted in the codebook, `season` is encoded as a numeric variable taking the values 1, 2, 3 or 4 to denote the seasons: 1: winter, 2: spring, 3: summer, 4: fall

Let's recode the `season` variable as a character variable with the name of the season as its value. After you have recoded `season`, you can check its distribution with  `count` in this code chunk.

```{r recode-season, exercise = TRUE, exercise.setup="glimpse-data-setup", exercise.completion=FALSE}

dcbikeshare <- ___ %>%
  mutate(
    season = case_when(
    ___
    )
  )

dcbikeshare %>% count(season) 
```


```{r recode-season-hint-1}
dcbikeshare <- dcbikeshare %>%
  mutate(
    season = case_when(
    season == 1 ~ "winter",
    ___
    )
  )
```

```{r recode-season-hint-2}
dcbikeshare <- dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    )
  )

```



You will notice that $R$ arranges the levels of `season` in alphabetical order. The first level, `fall` will be  automatically chosen as the baseline level in any model you fit with `season` in it. To change the baseline level, you can use `fct_relevel` inside `mutate` as follows:

```{r relevel-season-setup}
dcbikeshare <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ) )

```

```{r relevel-season, exercise=TRUE}

dcbikeshare <- dcbikeshare %>% 
              mutate( season = fct_relevel(season, "spring"))

dcbikeshare %>% count(season)
```

If everything worked as intended, you should see that `spring` has sprung to first place. 

### Recode `holiday` and `workingday`

As noted in the codebook, the variables `holiday` and `workingday` take the value of 0 and 1 to denote the labels "no" and "yes". 

Recode the variables `holiday` and `workingday` as character variables and replace the numeric values with their corresponding labels. 

Again, use the `count` function to examine your work. 

```{r recode-holiday-workingday-setup}
dcbikeshare <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ),
    season = fct_relevel(season, "spring")
  )
```

```{r recode-holiday-workingday, exercise = TRUE, message = FALSE, exercise.completion=FALSE}
dcbikeshare <- ___ %>%
  ___(
    ___
  )
```

```{r recode-holiday-workingday-hint-1}
Use mutate() the variables using ifelse()
```

```{r recode-holiday-workingday-hint-2}
dcbikeshare <- dcbikeshare %>%
  mutate(
    holiday = ifelse(___),      
    ___
  )
```

```{r recode-holiday-workingday-hint-3}
dcbikeshare <- dcbikeshare %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    workingday = ifelse(___)
  )
```

```{r recode-holiday-workingday-hint-4}
dcbikeshare <- dcbikeshare %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    workingday = ifelse(workingday == 0, "no", "yes"),
  )

dcbikeshare %>% count(holiday)

dcbikeshare %>% count(workingday) 
```



### Recode year

The variable `yr` is a numeric variable which takes the values 0 and 1 to denote the year the data is from. Let's recode `yr` to be a character variable with levels 2011 and 2012, with 2011 as the baseline level.


```{r recode-yr-setup}
dcbikeshare <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ),
    season = fct_relevel(season, "spring")) %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    holiday = fct_relevel(holiday, "no"),    
    workingday = ifelse(workingday == 0, "no", "yes"),
    workingday = fct_relevel(workingday, "no")
  )
```

```{r recode-yr, exercise = TRUE, message = FALSE, exercise.completion = FALSE}
dcbikeshare <- ___ %>%
  ___(
    ___
  )
```

```{r recode-yr-hint-1}
Take a look at the previous exercise for some inspiration!
```

```{r recode-yr-hint-2}
dcbikeshare <- dcbikeshare %>%
  mutate(
    yr = ___(___),
    yr = ___(___)
  )
```

```{r recode-yr-hint-3}
dcbikeshare <- dcbikeshare %>%
  mutate(
    yr = ifelse(___),
    yr = fct_relevel(___)
  )
```

```{r recode-yr-hint-4}
dcbikeshare <- dcbikeshare %>%
  mutate(
    yr = ifelse(yr == 0, "2011", "2012"),
    yr = fct_relevel(yr, "2011")
  )
```



### Recode `weathersit`

The last categorical variable we will consider is `weathersit`. This variable refers to the amount of precipitation (rain/snow).  Here is its description again from the codebook. 


| Variable name    | Description 
|:--------|:-------------------------------------------------------------
| `weathersit`| 1: Clear, Few clouds, Partly cloudy, Partly cloudy
|             | 2: Mist + Cloudy, Mist + Broken clouds, Mist + Few clouds, Mist
|             | 3: Light Snow, Light Rain + Thunderstorm + Scattered clouds, Light Rain + Scattered clouds
|             | 4: Heavy Rain + Ice Pallets + Thunderstorm + Mist, Snow + Fog


 Recode the `weathersit` variable as 1 - clear, 2 - mist, 3 - light precipitation, and 4 - heavy precipitation, with clear as the baseline. 


Again, don't forget to take a look at the distribution of the transformed variable with the `count` function.

```{r recode-weathersit-setup}
dcbikeshare <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ),
    season = fct_relevel(season, "spring")) %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    holiday = fct_relevel(holiday, "no"),    
    workingday = ifelse(workingday == 0, "no", "yes"),
    workingday = fct_relevel(workingday, "no")
  ) %>%
  mutate(
    yr = ifelse(yr == 0, "2011", "2012"),
    yr = fct_relevel(yr, "2011")
  )
```

```{r recode-weathersit, exercise = TRUE, warning=TRUE, message=FALSE, exercise.completion=FALSE}
dcbikeshare <- ___ %>%
  ___(
    ___
    ),
    ___
  )
```

```{r recode-weathersit-hint-1}
dcbikeshare <- dcbikeshare %>%
  mutate(
    weathersit = case_when(
      ___
    ),
    ___ = ___
  )
```

```{r recode-weathersit-hint-2}
dcbikeshare <- dcbikeshare %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      ___
    ),
    ___ = ___
  )
```

```{r recode-weathersit-hint-3}
dcbikeshare <- dcbikeshare %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      weathersit == 2 ~ "mist",
      weathersit == 3 ~ "light precipitation",
      weathersit == 4 ~ "heavy precipitation"
    ),
    ___ = ___
  )
```

```{r recode-weathersit-hint-4}
dcbikeshare <- dcbikeshare %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      weathersit == 2 ~ "mist",
      weathersit == 3 ~ "light precipitation",
      weathersit == 4 ~ "heavy precipitation"
    ),
    weathersit = fct_relevel(___)
  )
```

```{r recode-weathersit-hint-5}
dcbikeshare <- dcbikeshare %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      weathersit == 2 ~ "mist",
      weathersit == 3 ~ "light precipitation",
      weathersit == 4 ~ "heavy precipitation"
    ),
    weathersit = fct_relevel(weathersit, "clear")
  )
```


At this point, you might be wondering why the "heavy precipitation" category is missing from the distribution tabulated by `count`. 
This is simply because while it is possible for the `weathersit` variable to have the value 4 for heavy precipitation, this never actually occurs in the data.
Consequently, the `weathersit` factor only has three levels.


## Creating new variables

### Create raw values for temperature, humidity, and windspeed

In this section we turn our attention to the numeric variables shown below. 

| Variable name    | Description 
|:--------|:-------------------------------------------------------------
| `temp` 			| Normalized temperature in Celsius. The values are divided by 41 (max)
| `atemp`			| Normalized feeling temperature in Celsius. The values are divided by 50 (max)
| `hum` 			| Normalized humidity. The values are divided by 100 (max)
| `windspeed`	| Normalized wind speed. The values are divided by 67 (max)

As you can see, these variables have been normalized by the maximum value for the variable. Let's go ahead and 
calculate raw temperature, feeling temperature, humidity, and windspeed as their values given in the dataset multiplied by the maximum raw values stated in the codebook for each variable. 
Instead of writing over the existing variables, create new ones called `temperature_raw`, `feeling_temperature_raw`, `humidity_raw`, `windspeed_raw`.

```{r recode-measures-setup}
dcbikeshare <- dsbox::dcbikeshare %>%
  mutate(
    season = case_when(
      season == 1 ~ "winter",
      season == 2 ~ "spring",
      season == 3 ~ "summer",
      season == 4 ~ "fall"
    ),
    season = fct_relevel(season, "spring")) %>%
  mutate(
    holiday = ifelse(holiday == 0, "no", "yes"),      
    holiday = fct_relevel(holiday, "no"),    
    workingday = ifelse(workingday == 0, "no", "yes"),
    workingday = fct_relevel(workingday, "no")
  ) %>%
  mutate(
    yr = ifelse(yr == 0, "2011", "2012"),
    yr = fct_relevel(yr, "2011")
  ) %>%
  mutate(
    weathersit = case_when(
      weathersit == 1 ~ "clear",
      weathersit == 2 ~ "mist",
      weathersit == 3 ~ "light precipitation",
      weathersit == 4 ~ "heavy precipitation"
    ),
    weathersit = fct_relevel(weathersit, "clear")
  )
```

```{r recode-measures, exercise = TRUE, exercise.completion=FALSE}
dcbikeshare <- ___ %>%
  ___(
    ___
  )
```

```{r recode-measures-hint-1}
dcbikeshare <- ___ %>%
  mutate(
    ___= ___,
    ___
  )
```

```{r recode-measures-hint-2}
dcbikeshare <- ___ %>%
  mutate(
    temperature_raw = ___,
    ___
  )
```

```{r recode-measures-hint-3}
dcbikeshare <- ___ %>%
  mutate(
    temperature_raw = temp * 41,
    ___
  )
```

```{r recode-measures-hint-4}
dcbikeshare <- ___ %>%
  mutate(
    temperature_raw = temp * 41,
    feeling_temperature_raw = atemp * ___,
    humidity_raw = ___ * ___,
    windspeed_raw = ___ * ___
  )
```

```{r recode-measures-hint-5}
dcbikeshare <- dcbikeshare %>%
  mutate(
    temperature_raw = temp * 41,
    feeling_temperature_raw = atemp * 50,
    humidity_raw = hum * 100,
    windspeed_raw = windspeed * 67
  )
```



## Bike rentals and temperature


Fit a linear model predicting total daily bike rentals `cnt` from raw daily temperature variable you just created, and use `tidy`to print the output and answer the questions below.

```{r recreate-viz-setup}
dcbikeshare <- mutatedbike
```


```{r rentals-vs-temp, exercise = TRUE, exercise.setup = "recreate-viz-setup"}
cnt_tmp <- lm(___ ~ ___, data = ___)
cnt_tmp %>% tidy()
```


```{r rentals-vs-temp-hint-1}
cnt_tmp <- lm(cnt ~ ___, data = dcbikeshare)
cnt_tmp %>%
  tidy()
```

```{r rentals-vs-temp-hint-2}
cnt_tmp <- lm(cnt ~ temperature_raw, data = dcbikeshare)
cnt_tmp %>%
  tidy() 
```

```{r rentals-vs-temp-check}
	 grade_this({
   if(identical(floor(.result$estimate[2]), 161)) {
     pass("You have fit the model correctly")
   }
   if(identical(floor(.result$estimate[1]), 9)) {
     fail("Did you maybe try to predict the weather from the number of bike rentals?")
   }
   if(identical(floor(.result$estimate[1]),11)) {
     fail("Did you maybe try to predict the weather from the number of bike rentals?")
   }
   if(identical(floor(.result$estimate[1]), 0)) {
     fail("Did you maybe try to predict the weather from the number of bike rentals?")
   }
   fail("Not quite. Look at the hints for help!")
 })
``` 

Now, based on your findings, answer the following question. 

```{r rentals-vs-temp-quiz}
question("Which interpretations are correct?",
  answer("The model has the intercept at approximately 1215, which means that at a day with a temperature of 0 degrees celsius, we can except to have, on average, 1215 bike rentals.",
         correct = TRUE),
  answer("The model has the intercept at approximately 162, which means that at a day with a temperature of 0 degrees celsius, we can expect to have, on average, 162 bike rentals."),
  answer("The slope of the model is approximately 162, which means that, for each additional degree celsius, we can expect bike rentals to increase by 162.",
         correct = TRUE),
  answer("The slope of the model is approximately 1215, which means that, for each additional degree celsius, we can expect bike rentals to increase by 1215."),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

Use the chunk below to print out the $R^2$ of the model.

```{r rentals-vs-temp-r2-setup}
dcbikeshare <- mutatedbike
cnt_tmp <- lm(cnt ~ temperature_raw, data = dcbikeshare)
```

```{r rentals-vs-temp-r2, exercise = TRUE, exercise.setup="rentals-vs-temp-r2-setup"}
glance(___) %>% select(___)
```

```{r rentals-vs-temp-r2-solution}
glance(cnt_tmp) %>% select(r.squared)
```

```{r rentals-vs-temp-r2-check}
grade_this_code("Your solution is correct!")
```

Now, answer the following question:

```{r rentals-vs-temp-r2-quiz}
question("What does the R-squared mean in this context?",
  answer("39.4% of the variability in the number of total daily bike rental is explained by temperature.",
         correct = TRUE),
  answer("The model predicts the correct number of bike rentals 39.4% of the time"),
  answer("On average, the number of bike rentals predicted by the model differs from the actual number of bike rentals by 39.4%"),
  answer("On average, the number of bike rentals predicted squared by the model differs from the actual number of bike rentals squared by 39.4%"),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Bike rentals vs feeling temperature

Fit another linear model predicting total daily bike rentals `cnt` from `feeling_temperature_raw` you just created in the previous section. Print the estimates.
Then, proceed to answering the questions below.

```{r rentals-vs-atemp, exercise = TRUE, exercise.setup = "recreate-viz-setup"}
cnt_atmp <- lm(___ ~ ___, data = dcbikeshare)
___ %>%
  ___
```

```{r rentals-vs-atemp-hint-1}
Look at the previous question for help!
```

```{r rentals-vs-atemp-solution}
cnt_atmp <- lm(cnt ~ feeling_temperature_raw, data = dcbikeshare)
cnt_atmp %>%
  tidy()
```

```{r rentals-vs-atemp-check}
grade_this({
  if(identical(floor(.result$estimate[2]), 150)) {
    pass("You have fit the model correctly")
  }
  if(identical(floor(.result$estimate[1]), 11)) {
    fail("Did you maybe try to predict the weather from the number of bike rentals?")
  }
  if(identical(floor(.result$estimate[1]), 0)) {
    fail("Did you maybe try to predict the weather from the number of bike rentals?")
   }
  if(identical(floor(.result$estimate[1]), 1214)) {
    fail("Did you use normalized temperature instead of raw temperature?")
   }
  fail("Not quite. Look at the hints for help!")
})
``` 

Now, based on your findings, answer the following question. 

```{r rentals-vs-atemp-quiz}
question("Which interpretations are correct?",
  answer("The model has the intercept at approximately 946, which means that at a day with an apparent temperature of 0 degrees celsius, we can expect to have, on average, 946 bike rentals.",
         correct = TRUE),
  answer("The model has the intercept at approximately 150, which means that at a day with an apparent temperature of 0 degrees celsius, we can expect to have, on average, 150 bike rentals."),
  answer("The slope of the model is approximately 150, which means that, for each additional degree in apparent temperature, we can expect bike rentals to increase by 150.",
         correct = TRUE),
  answer("The slope of the model is approximately 946, which means that, for each additional degree in apparent temperature, we can expect bike rentals to increase by 946."),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

Use the chunk below to print out the $R^2$ of the model.

```{r rentals-vs-atemp-r2-setup}
dcbikeshare <- mutatedbike
cnt_atmp <- lm(cnt ~ feeling_temperature_raw, data = dcbikeshare)
```

```{r rentals-vs-atemp-r2, exercise = TRUE, exercise.setup="rentals-vs-atemp-r2-setup"}
glance(___) %>% select(___)
```

```{r rentals-vs-atemp-r2-solution}
glance(cnt_atmp) %>% select(r.squared)
```

```{r rentals-vs-atemp-r2-check}
grade_this_code("Your solution is correct!")
```

Is the following statement true or false?

```{r rentals-vs-atemp-r2-quiz}
question("39.8% of the variability in the number of total daily bike rentals is explained by apparent temperature.",
  answer("True",
         correct = TRUE),
  answer("False"),
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

Now, based on your findings from the two questions above, answer the following:

```{r better-predictor-quiz}
question("Which is the better explanatory variable of bike rentals?",
  answer("Feeling temperature is the better predictor, since the R-squared is slightly higher.",
         correct = TRUE),
  answer("Feeling temperature is the better predictor, since the R-squared is slightly lower.",
         message = "Determine the R-squared values again to find the correct answer!"),
  answer("Absolute temperature is a worse predictor, since the R-squared is higher.",
         message = "Determine the R-squared values again to find the correct answer!"),
  answer("Absolute temperature is a better predictor, because the R-squared is higher.",
         message = "Remember we want the higher R-squared value when choosing a model."),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Bike rentals, holiday and temperature

Fit a model predicting total daily bike rentals `cnt` from raw daily feeling temperature, whether the day is a `holiday` or not as well as the interaction between the 
feeling temperature and holiday.


```{r m-int-model, exercise = TRUE, exercise.setup = "recreate-viz-setup"}
cnt_tmp_hol <- lm(___ ~ ___, data = ___)
tidy(cnt_tmp_hol)
```


```{r m-int-model-hint-1}
cnt_tmp_hol <- lm(cnt ~ ___ + ___ + ___, 
                  data = dcbikeshare)

```

```{r m-int-model-hint-2}
cnt_tmp_hol <- lm(cnt ~ feeling_temperature_raw + ___ + ___, 
                  data = dcbikeshare)

```

```{r m-int-model-hint-3}
cnt_tmp_hol <- lm(cnt ~ feeling_temperature_raw + holiday + ___, 
                  data = dcbikeshare)

```


```{r m-int-model-solution}
cnt_tmp_hol <- lm(cnt ~ feeling_temperature_raw + holiday + holiday:feeling_temperature_raw, data = dcbikeshare)
tidy(cnt_tmp_hol)
```

```{r m-int-model-check}
grade_this({
  if(identical(floor(.result$estimate[1]), 1002)) {
    pass("You have fit the model correctly")
  }
  if(identical(floor(.result$estimate[1]), 970)) {
    fail("Did you forget to model the interaction effect between feeling temperature and holiday?")
  }
   fail("Not quite. Look at the hints for help!")
})
``` 

Investigate how the model values change depending on whether a day is a holiday or not and answer the following questions.

```{r temp-slopes-quiz}
question("Which of the following statements are correct?",
  answer("The intercept for non-holidays is higher than for holidays",
         correct = TRUE),
  answer("The intercept for non-holidays is lower than for holidays",
         message = "Look at the regression output of the model for a hint!"),
  answer("The slope for feeling temperature is steeper for non-holidays than for holidays.",
         message = "Look at the end of the regression output of for a hint!"),
  answer("The slope for feeling temperature is steeper for holidays than for non-holidays.",
         correct = TRUE),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```



### Full model

Now we will go ahead and fit the kitchen sink model: one with everything in it. 
Fit a model predicting total daily bike rentals `cnt` from `season`, `yr`, whether the day is `holiday` or not, whether the day is a `workingday` or not, the `weathersit` category, `temperature_raw`, `feeling_temperature_raw`, `humidity_raw`, and `windspeed_raw`, as well as the interaction between feeling temperature and holiday.

```{r m-full, exercise = TRUE, exercise.setup = "recreate-viz-setup"}
cnt_full <- lm(___ ~ ___, data = ___)
tidy(cnt_full)
```

```{r m-full-hint-1}
cnt_full <- lm(cnt ~ ___ + feeling_temperature_raw * holiday, data = ___)
tidy(cnt_full)
```

```{r m-full-solution}
cnt_full <- lm(cnt ~ season + yr + holiday + workingday + weathersit +
        temperature_raw + feeling_temperature_raw + humidity_raw +
        windspeed_raw + feeling_temperature_raw:holiday, 
      data = dcbikeshare)
tidy(cnt_full)
```

```{r m-full-check}
grade_this({
  if(identical(floor(.result$estimate[1]), 2715)) {
    pass("You have fit the model correctly")
  }
  if(identical(floor(.result$estimate[1]), 2688)) {
    fail("Did you forget to model the interaction effect between feeling temperature and holiday?")
  }
   fail("Not quite. Look at the hints for help!")
})
``` 

```{r adj-r-squared-setup}
dcbikeshare <- mutatedbike
cnt_full <- lm(cnt ~ season + yr + holiday + workingday + weathersit +
        temperature_raw + feeling_temperature_raw + humidity_raw +
        windspeed_raw + feeling_temperature_raw:holiday, 
      data = dcbikeshare)
```

Record adjusted $R^2$ of the model.

```{r adj-r-squared, exercise = TRUE, exercise.setup="adj-r-squared-setup"}

```

```{r adj-r-squared-hint-1}
glance(cnt_full) %>% select(adj.r.squared)
```

```{r adj-r-squared-solution}
glance(cnt_full) %>% select(adj.r.squared)
```


```{r adj-r-squared-check}
grade_this_code("Your solution is correct!")
```


According to this model, assuming everything else is the same, in which season does the model predict total daily bike rentals to be highest and which to be the lowest?

```{r intercept-quiz}
question("Select the correct answer",
  answer("Highest: Fall, Lowest: Winter",
         correct = TRUE),
  answer("Highest: Summer, Lowest: Winter"),
  answer("Highest: Fall, Lowest: Spring"),
  answer("Highest: Spring, Lowest: Winter"),
  answer("Highest: Spring, Lowest: Fall"),
  answer("Highest: Summer, Lowest: Fall"),
  correct = "Correct!",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

<!---
### Backward selection

The full regression model from the previous section may have variables in it that do not necessarily add much once the other variables are included. This leads to needlessly complex models. 

In this section we will perform the first step of backward selection by fitting a series of models, each with one explanatory variable removed from the full model you fit in the previous exercise and record the adjusted $R^2$s of each of these models.

In the chunk below, fit the full model without the `season` variable and print out the $R^2$ score. An easy way to drop one variable from the model is to use the `update` function. 

For example, to drop the interaction between `holiday` and `feeling_temperature_raw` from the model `cnt_tmp_hol`, we would specify the update function with the model object and a new formula. The new formula can use the period as short hand for keep everything on either the left or right hand side of the formula and the plus or minus sign used to add or remove terms to the model.

```{r echo=T, eval=F}
update(cnt_tmp_hol, . ~ . - holiday:feeling_temperature_raw)

```


```{r remove-season, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}
rm_season <- update( cnt_full, ___ )
glance(rm_season)$adj.r.squared
```

```{r remove-season-solution}
rm_season <- update(cnt_full, .~. - season )
glance(rm_season)$adj.r.squared
```




```{r remove-season-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 4)), 0.7491)) {
     pass("You have fit the model correctly")
   }
   fail("Not quite. Look at the hint for help and remember to just remove the season variable from the full model!")
})
``` 

In the chunk below, fit the full model without the `year` variable and print out the adjusted $R^2$ score.

```{r remove-year, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}
rm_year <- update( ___, ___) 
glance(___)$___
```

```{r remove-year-solution}
rm_year <- update( cnt_full, .~. - yr) 
glance(rm_year)$adj.r.squared
```

```{r remove-year-check}
grade_this({
  if(identical(as.numeric(round(.result, digits = 5)), 0.5536)) {
    pass("You have fit the model correctly")
  }
   fail("Not quite. Remember to just remove the year variable from the full model!")
})
``` 

In the chunk below, fit the full model without the `holiday` variable and print out the adjusted $R^2$ score.
Remember to also remove the interaction effect between `holiday` and `feeling_temperature_raw`.

```{r remove-holiday, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}

```

```{r remove-holiday-solution}
rm_holiday <- update( cnt_full, .~. - holiday - holiday:feeling_temperature_raw )
glance(rm_holiday)$adj.r.squared
```

```{r remove-holiday-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 4)), 0.8173)) {
     pass("You have fit the model correctly")
   }
   if(identical(as.numeric(round(.result, digits = 3)), 0.820)) {
     fail("Did you forget to remove the interaction effect?")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Fit the full model without the `workingday` variable using the code chunk below! Remember to record the adjusted $R^2$ score.

```{r remove-workingday, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}

```


```{r remove-workingday-solution}
rm_workingday<- update( cnt_full, .~. - workingday )
glance(rm_workingday)$adj.r.squared
```


```{r remove-workingday-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 4)), 0.8196)) {
     pass("You have fit the model correctly")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Fit the full model without the `weathersit` variable in the chunk below and record the adjusted $R^2$ score.

```{r remove-weathersit, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}

```

```{r remove-weathersit-solution}
rm_weathersit<- update( cnt_full, .~. - weathersit )
glance(rm_weathersit)$adj.r.squared
```

```{r remove-weathersit-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 5)), 0.79847)) {
     pass("You have fit the model correctly")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Remove the `temperature_raw` variable in the chunk below and record the adjusted $R^2$ score.

```{r remove-tempraw, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}


```

```{r remove-tempraw-solution}
rm_temp_raw<- update( cnt_full, .~.-temperature_raw )
glance(rm_temp_raw)$adj.r.squared
```

```{r remove-tempraw-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 5)), 0.81813)) {
     pass("You have fit the model correctly")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

In the chunk below, fit the full model without the `feeling_temperature_raw` variable and record the adjusted $R^2$ score. Don't forget to remove the interaction of `holiday` and `feeling_temperature_raw`.


```{r remove-feeltempraw, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}


```

```{r remove-feeltempraw-solution}
rm_feeltemp_raw <- update( cnt_full, .~. - feeling_temperature_raw - feeling_temperature_raw:holiday )
glance(rm_feeltemp_raw)$adj.r.squared
```

```{r remove-feeltempraw-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 5)), 0.81988)) {
     pass("You have fit the model correctly")
   }
   if(identical(as.numeric(round(.result, digits = 3)), 0.820)) {
     fail("Did you forget to remove the interaction effect?")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Fit the full model without the `humidity_raw` variable and record the adjusted $R^2$ score.

```{r remove-humidityraw, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}


```

```{r remove-humidityraw-solution}
rm_humidity_raw <- update(cnt_full,  .~. - humidity_raw)
glance(rm_humidity_raw)$adj.r.squared
```

```{r remove-humidityraw-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 4)), 0.8151)) {
     pass("You have fit the model correctly")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Fit the full model without the `windspeed_raw` variable and record the adjusted $R^2$ score.

```{r remove-windraw, exercise = TRUE, exercise.setup = "adj-r-squared-setup"}


```

```{r remove-windraw-solution}
rm_windspeed_raw <- update(cnt_full , .~. - windspeed_raw)
glance(rm_windspeed_raw)$adj.r.squared
```

```{r remove-windraw-check}
grade_this({
  if(identical(as.numeric(round(.result, digits = 4)), 0.8106)) {
    pass("You have fit the model correctly")
  }
  fail("Not quite. Did you use the full model as your starting point?")
})
```

Last, remove the interaction effect between `holiday` and `feeling_temperature_raw` without removing the two variables individually.
Then, record the adjusted $R^2$.

```{r remove-interaction, exercise = TRUE, exercise.setup = "adj-r-squared"}


```

```{r remove-interaction-solution}
rm_interaction <- update( cnt_full, .~.-holiday:feeling_temperature_raw )
glance(rm_interaction)$adj.r.squared
```

```{r remove-interaction-check}
grade_this({
   if(identical(as.numeric(round(.result, digits = 4)), 0.8197)) {
     pass("You have fit the model correctly")
   }
   if(identical(as.numeric(round(.result, digits = 4)), 0.8174)) {
     fail("Did you also remove the two variables individually in addition to removing the interaction effect?")
   }
   fail("Not quite. Did you use the full model as your starting point?")
})
```

Now, using your findings from the previous exercises, answer the following question: 

```{r backward-selection-quiz}
question("Based on the above results, which of these variables should you definitely not drop from the model of these models, if any, gives the highest improvement over the full model?",
  answer("rm_season, rm_workingday, rm_year"),
  answer("rm_season, rm_interaction, rm_holiday"),
  answer("rm_winspeed_raw, rm_humidity_raw, rm_year"),
  answer("rm_feeltemp_raw, rm_workingday, rm_interaction"),
  answer("All"),
  answer("None", correct = TRUE),
  correct = "Correct!",
  incorrect = "Look up the adjusted R-squared of the full models and compare it with the others. Remember we want the highest possible adjusted R-squared score.",
  allow_retry = TRUE
)
```

--->

## Wrap up

Congratulations!

Hopefully, you have enjoyed this lesson on modelling and data wrangling.

If this is a participation activity, generate the hash and submit. 

```{r context="server"}
learnrhash::encoder_logic(strip_output=TRUE)
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(ui_after=NULL, ui_before=default_ui(url="https://canvas.uw.edu/courses/1483091") )
```



## Acknowledgment

This tutorial is excerpted from Data Science in a Box.