---
title: "Road Traffic Accidents"
output: 
     learnr::tutorial:
     progressive: true
     allow_skip: true
     css: "css/font-size.css"
runtime: shiny_prerendered
description: >
  Get more practice with data visualization and wrangling in the tidyverse: 
  practice visualization for categorical variables with ggplot, tabyl and mosaicplot 

---


```{r setup, include=FALSE}

#load packages ------------------------------------------
library(learnr)
library(tidyverse)
library(janitor)
#library(gradethis)

# set options for exercises and checking ---------------------------------------
#gradethis_setup()

# hide non-exercise code chunks ------------------------------------------------
knitr::opts_chunk$set(echo = FALSE)

# read in data ----------------------------------
accidents<-read_csv("data/accidents_2020.csv",skip=1,
                    col_names=c("id", "addr_type","location","severity_code","severity_desc","collision_type","person_count","ped_count","pedcyl_count","veh_count","injuries","serious_injuries","fatalities","inc_date","under_infl","weather","road","light","speeding","hit_parked_car"))


accidents <- accidents %>%  mutate(light=na_if(light,"Unknown"), 
                                  light=na_if(light, "Other") )
```

---

## Introduction

In this tutorial, we will look at traffic accident data in Seattle.
The data are open and made available [online](https://data-seattlecitygis.opendata.arcgis.com/datasets/5b5c745e0f1f48e7a53acec63a0022ab_0) by Seattle GeoData. 
It covers all recorded accidents in Seattle from Jan 1 2020 through Dec 16 2020; some of the variables were deleted for the purposes of this exercise.

```{r fig, out.width="100%"}

knitr::include_graphics("images/traffic.jpg")

```

### Learning goals

- Practice data visualization for categorical data  with ggplot. 
- Make contingency tables for two categorical variables using tabyl.
- Create and interpret mosaic plots.


### Packages

We will use the **tidyverse** package for the analysis and the **janitor** package for creating contingency tables.
These packages are already installed for you, so you load it as usual with the `library` function:

```{r load-packages, exercise = TRUE}


```

```{r load-packages-solution}
library(tidyverse)
library(janitor)
```

## Data and codebook

The data has been uploaded with this tutorial and is called `accidents` .

Below is an excerpt of the data dictionary also known as the code book.  Note that it is long (there are lots of variables in the data), but we will be using a limited set of the variables in this tutorial.

| Header                        |  Description
|:----------------|:--------------------------------
| `id` | unique identifier
| `addr_type` | Collision address type: **Alley**, **Block**, **Intersection** 
| `location` | Text description of the general location of the collision
| `severity_code` |3 - fatality, 2b - serious injury, 2 -injury, 1 - property damage, 0 - unknown
| `severity_desc   ` | Description of severity: **Fatality Collision**, **Injury Collision**, **Property Damage Only Collision**, **Serious Injury Collision**, **Unknown**
| `collision_type` | Collision type: **Angles**, **Cycles**, **Head On**, **Left Turn**, **Parked Car**, **Pedestrian**, **Rear Ended**, **Right Turn**, **Sideswipe**
| `person_count` | Number of people involved in the collision
|`ped_count` |Number of pedestrians involved in the collision
| `pedcyl_count` | Number of bicycles involved in the collision 
|`veh_count`| Number of vehicles involved in the collision (0, 1, 2, etc)
|`injuries`|Number of total injuries in the collision
|`serious_injuries`|Number of serious injuries in the collision
|`fatalities`| Number of fatalities in the collision
|`inc_date`| Date of incident
|`under_infl`| Whether or not the driver was under the influence of drugs or alcohol: **Y**, **N**
|`weather`|A text description of the weather conditions during the collision
|`road`|The condition of the road during the collision: **Dry**, **Wet**, **Ice**, **Sand/Mud/Dirt**, **Snow/Slush**, **Standing Water**, **Unknown**
|`light`|Light conditions during the collision: **Dark - No Street Lights**, **Dark - Street Lights Off**, **Dark - Street Lights On**, **Dark - Unknown Lighting**, **Dusk**, **Dawn**, **Daylight**
|`speeding`|  Whether or not speeding was a factor in the collision: **Y**, **N**
|`hit_parked_car`|Whether or not the collision involved hitting a parked car: **Y**, **N**

### First look at the data

You can take a peek at the data using the `glimpse()` function in the box below. 

```{r glimpse-data, exercise = TRUE}


```

```{r quiz-rows}
question("What does each row in the dataset represent?",
    answer("The registration number of a car"),
    answer("The location of an accident"),
    answer("A recorded accident",
           correct = TRUE,
           message = "Each row in the dataset contains all information relating to an individual recorded accident in Seattle."),
    answer("An insurance claim "),
    allow_retry = TRUE
  )
```

How many accidents were recorded in Seattle in 2020? The output from `glimpse` should give you the answer. But you can also use `nrow()` to verify the answer. Run it in the code chunk below. 

```{r nrow,exercise=TRUE}

```

```{r nrow-hint-1}
Each row represents one recorded accident!
```

```{r nrow-hint-2}
Try using nrow()!
```

<!---
```{r nrow-check}
grade_this({
  if(identical(.result, 6509) | identical(.result, 6509L)) {
    pass("Great! There are 6,509 rows, therefore, 6,509 accidents were recorded in Seattle in 2020.")
    }
  if(identical(.result, 20) | identical(.result, 20L)) {
    fail("Each observation is represented in one row. Did you calculate the number of columns instead of rows?")
    }
  fail("Not quite. Each observation is represented in one row. Try looking at the hints for some help!")
})
```
--->

How many variables are recorded on each crash? Again, the output from `glimpse` gives the answer. But you can use `ncol()` to verify the answer. Try it.

```{r ncol, exercise=TRUE}

```

```{r ncol-hint-1}
Each variable is encoded is displayed as a column.
```

```{r ncol-hint-2}
Try using ncol()!
```

<!---
```{r ncol-check}
grade_this({
  if(identical(.result, 20) | identical(.result, 20L)) {
    pass("Great! Since there are 20 columns in the dataset, we know that 20 variables are recorded.")
  }
  if(identical(.result, 6509) | identical(.result, 6509L)) {
    fail("Each variable is recorded in a column. Did you maybe use the number of rows?")
  }
  fail("Not quite. Each variable is represented in a column. Try looking at the hints for some help!")
})
```

--->

```{r quiz-vars}
question("What type of variable is addr_type?",
    answer("Numerical"),
    answer("Categorical with ordered categories"),
    answer("Categorical with unordered categories ", correct=TRUE),
    allow_retry = TRUE
  )
```

## Multi-vehicle accidents

How many accidents involved 2 or more vehicles?  Filter on `veh_count` to find out. 

``` {r filter-multi-car, exercise = TRUE}

```

```{r filter-multi-car-hint-1}
Use filter() to find the rows that match the criteria.
```

```{r filter-multi-car-hint-2}
Review the data dictionary, specifically the variable veh_count.
```

``` {r filter-multi-car-hint-3}
accidents %>%
  filter(veh_count >= ___ ) %>% 
  nrow()
```

``` {r filter-multi-car-hint-4}
accidents %>%
  filter(veh_count >= 2) %>% 
  nrow()
```

<!---
```{r filter-multi-car-check}
grade_this({
  if(identical(.result, 3673) | identical(.result, 3673L)) {
    pass("Correct! There are 3,673 rows that correspond to accidents with 2 or more vehicles involved.")
  }
  fail("Take a peek at the hints!")
})
```
--->

How many accidents involved `speeding`?  

``` {r filter-speeding, exercise = TRUE}


```

```{r filter-speeding-hint-1}
Review the data dictionary, specifically the variable speeding
```

```{r filter-speeding-hint-2}
accidents %>%
    filter(speeding=="Y") %>%
    nrow()
```

<!---
```{r filter-speeding-check}
grade_this({
  if(identical(.result, 292) | identical(.result, 292L)) {
    pass("Nice work! There are 292 rows that correspond to accidents with speeding involved.")
  }
  fail("Take a look at the solution!")
})
```
--->

How many accidents involved 2 or more vehicles and speeding?

```{r filter-multi-car-speed, exercise=TRUE} 


```


```{r filter-multi-car-speed-hint-1}
accidents %>%
    filter(veh_count >= ___, speeding == ____ ) %>%
    nrow()
```

<!---
```{r filter-multi-car-speed-check}
grade_this({
  if(identical(.result, 164) | identical(.result, 164L)) {
    pass("There are 164 rows that correspond to accidents with two or more vehicles and speeding involved.")
  }
  fail("Take a look at the solution!")
})
```

--->

## Location of accident

The variable `addr_type` describes the location of the accident.
Where do the most accidents occur? Create a frequency table of the `addr_type` in decending order of frequency. See hints for help.


```{r location-frequency,exercise=TRUE}

```


```{r location-frequency-hint-1}
See the help for the `count()` function, specifically the 
`sort` argument for reporting the frequency table in descending order of counts, 
i.e. highest on top.
```

```{r location-frequency-hint-2}
accidents %>% 
  ___ %>% (___, sort=TRUE)
```

```{r location-frequency-hint-3}
accidents %>%
  count(addr_type, sort = TRUE)
```

<!---

```{r location-frequency-check}
grade_this({
  if(identical(.result$n[1], 4060L)) {
    pass("You have created the correct frequency table!")}
  fail("Not quite. See the hints for help!")
 })
```

--->


## Accident Severity

Suppose we wish to compare accident severity (`severity_desc`) by the location of the accident (`addr_type`). Let's recreate the following segmented bar plot with the values of `addr_type` on the horizontal axis. Note that the rows that are missing the address type are not represented in the bar plot. This is achieved by filtering the data using the `is.na` function with `addr_type` as the argument.    

```{r example-plot, fig.width=6,fig.height=4, echo=FALSE}

ggplot(data=accidents %>% filter(!is.na(addr_type)), 
  mapping=aes(x=addr_type, fill=severity_desc)) +
      geom_bar(position = "fill")+
      labs(x="Location of accident",
           y="Proportion",
           fill="Severity Description",
           title="Location and severity of accident")

```

```{r plot-severity, exercise = TRUE, fig.width=6, echo=FALSE, fig.height=4}

ggplot(data = ___,   mapping=aes(x = ___, fill = ___)) +
  geom___(___) +
  ___(x = ___, y = ___,
       ___ = ___, 
       title = ___) 
```

```{r plot-severity-hint-1}
#code out loud:
# - start with the accidents data frame
# - filter the data using !is.na(addr_type) 
# - map addr type to the x-axis and severity desc to the fill 
# - represent counts with a bar and 
# - choose position as fill to create segmented bars
# - give the plot a title, subtitle, axis labels 


```


```{r plot-severity-hint-2}
ggplot(data = accidents %>% filter( !is.na(addr_type) ), 
       mapping=aes(x = ___, fill = ___)) +
  geom_bar(___) +
  labs(x = ___, y = ___,
       fill = ___, 
       title = ___) 
```

```{r plot-severity-hint-3}
ggplot(data = accidents %>% filter( !is.na(addr_type) ), 
       mapping=aes(x = ___, fill = ___)) +
  geom_bar(position = ___) +
  labs(x = ___, y = ___,
       fill = "Severity Description", 
       title = ___) 
```

``` {r plot-severity-solution}
ggplot(data = accidents %>% filter( !is.na(addr_type) ), 
       mapping=aes(x = addr_type, fill = severity_desc)) +
  geom_bar(position = "fill") +
  labs(x = "Location of accident",
       fill = "Severity Description", 
       title = "Location and severity of accident")
```

## Light conditions

Let's now examine the relationship between the light conditions and whether or not it is related to parked cars being hit. In this section, our focus will be creating a contingency table and visualizing the relationship via a mosaicplot. 

But first, let's take a look at the `light` variable and its frequency distribution.

```{r light-values, exercise=TRUE}

accidents %>% count(light)
```

As we can see, there are four categories of darkness. Perhaps we can combine them into fewer categories in some sensible manner? Additionally, the category labels are pretty long and which can result in messy labels in certain types of plots, so at the very least we want to shorten the labels. 

### Recoding light

There are many ways to recode the levels of a categorical variable. In this tutorial, we will focus on using `mutate` and `case_when` to make the change. The `case_when` function from `dplyr` is essentially multiple `ifelse` statements combined together. Each argument within the `case_when` function is a two sided formula separated by a tilda symbol `~`. The left hand side of the formula determines which values match this case. The right hand side provides the replacement value. The last line of the `case_when` holds true for all the categories not mentioned in the previous statements and these are all set to the values in the original variable.



```{r, echo=TRUE}


accidents <- accidents %>% 
  mutate( light_short = case_when( light %in% c("Dark - No Street Lights",  "Dark - Street Lights Off") ~ "Dark", 
                                   light == "Dark - Street Lights On"                                    ~ "Dark, lit", 
                                                                      TRUE                               ~ light),
          light_short = na_if(light_short, "Dark - Unknown Lighting") )


```

What does the above code do? The `case_when` function creates a new variable called `light_short` by 

 - combining the categories "Dark - No Street Lights" and "Dark - Street Lights Off" and calling the new category "Dark" 

 - and shortening the label for the category "Dark - Street Lights On" by replacing it with "Dark, lit".  

The remaining categories are not altered by the `case_when` function. 

Within the same `mutate` function, we are then replacing the category
"Dark - Unknown Lighting" as a missing value  using the `na_if` function. This makes sense since we are interested in the impact of lighting which is unknown. 


Run the following to examine the frequency distribution of the new variable we have just created. Compare the distribution with that of the original variable at the head of this subsection. 

```{r change-light}

accidents <- accidents %>% 
  mutate( light_short = case_when( light %in% c("Dark - No Street Lights",  "Dark - Street Lights Off") ~ "Dark", 
                                   light %in% c("Dark - Street Lights On", "Dark - Unknown Lighting") ~ "Dark, lit", 
                                   TRUE                              
                           ~ light)  , 
         light_short = na_if(light_short, "Unknown"),
         light_short = na_if(light_short, "Dark - Unknown Lighting"))


```

```{r count-light-short, exercise=TRUE, exercise.setup="change-light"}

accidents %>% count(light_short)

```

Let's use this new more concise description of the light conditions to examine the relationship between light conditions and hitting parked cars.

### Light condition and hitting parked cars

Relationships between categorical variables are summarized via contingency tables. We can easily create a table using the `tabyl` function from the **janitor** package. Complete the blanks in the `tabyl` function to create a contingency table of the relationship between `light_short` and `hit_parked_car`.  

```{r light-hit-table, exercise=TRUE, exercise.setup="change-light"}

accidents %>%
          tabyl(___, ____, show_na=___) %>%
             adorn_totals()%>%
             adorn_percentages(denominator="row") %>%
             adorn_pct_formatting(digits=2) %>%
             adorn_ns()

```

```{r light-hit-table-hint}
#the explanatory goes first followed by the response
```



```{r table-interpretation}
question("If the variables are independent, what fraction of parked cars should be hit regardless of the light condition?", 
         answer(" 11.99% ",
                correct = TRUE, message="Good job!"),
         answer("27.45%%",
                message = "This is the fraction of parked cars that get hit in the dark"),
         answer("88.01%",
                message = "This is the fraction of parked cars that should not get hit if the variables are independent"),
        allow_retry = TRUE
         )
```

Comparing the percentage of hits that are expected under independence with the observed percentages in the contingency table confirms that the variables `light_short` and `hit_parked_car` are indeed related. In particular, accidents that occur during the daylight have much fewer instances of parked cars getting hit than would be expected if the variables were independent. Conversely, accidents that occur in the dark have many more instances of parked cars getting hit than should be expected. 

A mosaic plot is a visual tool for examining the relationship between two categorical variables. 

Create a mosaic plot of the relationship between `light_short` and `hit_parked_car` by filling in the blanks in the code below.


```{r light-severity-mosaic, exercise=TRUE, exercise.setup="change-light"}

mosaicplot(___ ~ ____, data = _____, shade = ____ )

```


```{r light-severity-mosaic-hint-1}

mosaicplot( __ ~ ___, data = ___, shade = TRUE)
```

```{r light-severity-mosaic-hint-2}

mosaicplot(hit_parked_car ~ light_short, data=accidents, shade=TRUE) 
```


Note that the mosaicplot function ignores the accidents for which the light conditions were missing by default, so we don't need to manually control how missing values are treated. 

Looking at the plot, do we see any reds and blues? These are the cells that depart from independence by having too many or too few cases in them. Compare the colored cells in the mosaicplot with the percentages in the contingency table. 

And finally, what can we conclude from the relationship between light conditions and hitting parked cars? Can we say that the darkness causes people to crash into parked cars? Why or why not? 

<!---
```{r mosaic-plot-interpretation}
question("Which of the following is false? Check all that apply.", 
         answer("Most accidents do not involve parked cars getting hit",
                message=""),
         answer("The variables are independent ", correct=TRUE,
                message = "Some of the cells are colored blue or red, which suggests an association"),
         answer("There is a higher chance of a parked car being hit in the dark than during day light hours",
                correct = TRUE),
        allow_retry = TRUE
         )
```

--->

## Wrap Up

Congratulations! You made it through the tutorial on visualizing categorical data! Hopefully you found it interesting and are feeling more comfortable with the concepts we have been learning in class today!

## Acknowledgements

The R code for creating this learnr tutorial was greatly informed by the learnr tutorials in [Data Science in a Box](https://datasciencebox.org/interactive-tutorials.html).


